Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name, email and UFID.
Please do not modify instruction cells or any cells with automated tests (marked with `[ASSERTS]`). Note: you can add new cells if you need them, but answers must be in the cells with `YOUR CODE HERE` or "YOUR ANSWER HERE" comments.

---

# Homework 0: Setting Up Your Environment!

## Preamble: Write your Name, Email and UFID

<!-- grade_id: name-email-ufid, points: 0, grade: false, solution: true -->
```python
NAME = 'Tikaharu Sharma'
EMAIL = 'tikaharusharma@ufl.edu'
UFID = 73535666

# YOUR CODE HERE
# raise NotImplementedError()
```

<!-- grade_id: name-email-ufid-asserts, points: 0, grade: true, solution: false -->
```python
""" [ASSERTS] Check that your name, email, and UFID is filled in."""
assert NAME != '' and NAME != 'Your name here.' and len(NAME) > 3
assert EMAIL != '' and EMAIL != 'Your email here.' and len(EMAIL) > 7
assert type(UFID) == int and UFID != 12345678 and UFID >= 10000000 and UFID <= 99999999
```

## Academic Integrity

### <span style="color:red;">This is an individual assignment. Academic integrity violations (i.e., cheating, plagiarism) will be reported to SCCR!</span><br/>
#### The official CISE policy recommended for such offenses is a course grade of E. Additional sanctions may be imposed by SCCR such as marks on your permanent educational transcripts, dismissal or expulsion.
#### Reminder of the Honor Pledge: On all work submitted for credit by Students at the University of Florida, the following pledge is either required or implied: *"On my honor, I have neither given nor received unauthorized aid in doing this assignment."*

#### Acknowledgement: Do you acknowledge and understand the academic integrity warning above? 

<!-- grade_id: academic-integrity, points: 0, grade: false, solution: true -->
```python
academic_integrity_acknowledgement = True
# YOUR CODE HERE
# raise NotImplementedError()
```

<!-- grade_id: academic-integrity-assert, points: 0, grade: true, solution: false -->
```python
""" [ASSERTS] Check that you acknowledge the academic integrity warning, you understanding it and have been reminded of the UF Honor Pledge."""
assert academic_integrity_acknowledgement
```

---
# [Task 1] Check that the cell below runs with no errors and that the version of Python is at least 3.10, the version of Pytorch is at least 2.2.0!

```python
# Load packages we need
import sys
import os

import numpy as np
import scipy as sp
import pandas as pd

import sklearn

import matplotlib 

import torch

# Let's check our software versions
print('### Python version: ' + sys.version)
print('### Numpy version: ' + np.__version__)
print('### Scipy version: ' + sp.__version__)
print('### Pandas version: ' + pd.__version__)
print('### Scikit-learn version: ' + sklearn.__version__)
print('### Matplotlib version: ' + matplotlib.__version__)
print('### Pytorch version: ' + torch.__version__)
print('------------')
```

Output:
```
### Python version: 3.10.19 (main, Oct 21 2025, 16:37:10) [Clang 20.1.8 ]
### Numpy version: 2.2.5
### Scipy version: 1.15.3
### Pandas version: 2.3.3
### Scikit-learn version: 1.7.2
### Matplotlib version: 3.10.8
### Pytorch version: 2.9.1
------------
```

### [Task 1] (15 Points) Paste in the cell below the output of the cell above (with the library versions)!

<!-- grade_id: task1-student-answer-paste, points: 0, grade: false, solution: true -->
```python
# Paste the output here:
### Python version: 3.10.19 (main, Oct 21 2025, 16:37:10) [Clang 20.1.8 ]
### Numpy version: 2.2.5
### Scipy version: 1.15.3
### Pandas version: 2.3.3
### Scikit-learn version: 1.7.2
### Matplotlib version: 3.10.8
### Pytorch version: 2.9.1
# YOUR CODE HERE
# raise NotImplementedError()
```

<!-- grade_id: task1-paste-tests, points: 15, grade: true, solution: false -->
```python
""" [ASSERTS] Check environment and versions."""
python_ver_split = sys.version.split('.')
assert int(python_ver_split[0]) >= 3 and int(python_ver_split[1]) >= 8
pt_ver_split = torch.__version__.split('.')
assert int(pt_ver_split[0]) >= 2 and int(pt_ver_split[1]) >= 2
```

---
# [Task 2] (10 Points) Tell us a little bit about yourself and what you hope to learn in this course

## [Task 2a] (5 Points) Write a brief description of your academic history (2-3 sentences is fine)

<!-- grade_id: task2a-input, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
# raise NotImplementedError()
print("I am a current Master's in Computer and Information Science student in the University of Florida. I did my undergraduate from Tribhuvan University, Nepal with a degree in Computer Engineering.")
```

Output:
```
I am a current Master's in Computer and Information Science student in the University of Florida. I did my undergraduate from Tribhuvan University, Nepal with a degree in Computer Engineering.
```

## [Task 2b] (5 Points) Briefly describe what you hope to learn in this course (a few sentences is fine)

<!-- grade_id: task2b-input, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
# raise NotImplementedError()
print("I enrolled in this class to gain a proper concept in the field of Machine Learning that will be very relevant to the work currently done in the industry. I have studied a lot of theory and concepts in the field of Machine Learning but have very limited practical exposure. So, I am expecting that I get a lot of opportunity in this class to gain hands on knowledge in Machine Learning by writing codes and building projects, that will ultimately help me in my industrial career.")
```

Output:
```
I enrolled in this class to gain a proper concept in the field of Machine Learning that will be very relevant to the work currently done in the industry. I have studied a lot of theory and concepts in the field of Machine Learning but have very limited practical exposure. So, I am expecting that I get a lot of opportunity in this class to gain hands on knowledge in Machine Learning by writing codes and building projects, that will ultimately help me in my industrial career.
```

---
# [Task 3] (25 Points) Familiarize yourself with Python and NumPy and write some code to demonstrate your understanding. (Note: the basics notebook is a great resource!)

## [Task 3a] (5 Points) Create a numpy array 'a' that contains the integers 0 to 10. The shape of it should be (11,).

<!-- grade_id: task3a-input, points: 0, grade: false, solution: true -->
```python
"""Put your code to create 'a' here.
"""
# YOUR CODE HERE
# raise NotImplementedError()
a = np.array([i for i in range(11)])
```

<!-- grade_id: task3a-test, points: 5, grade: true, solution: false -->
```python
""" [ASSERTS] Check the answer."""
print(a) # this should print: "[ 0 1  2  3  4  5  6  7  8  9 10]"
assert a.shape == (11,)
```

Output:
```
[ 0  1  2  3  4  5  6  7  8  9 10]
```

## [Task 3b] (5 Points) Let 'b' be the numpy array where every element is the square of the corresponding element in 'a'. The shape should be (11,).

<!-- grade_id: task3b-input, points: 0, grade: false, solution: true -->
```python
"""Put your code to create 'b' here.
"""
# YOUR CODE HERE
# raise NotImplementedError()
b = np.array([num**2 for num in a])
```

<!-- grade_id: task3b-test, points: 5, grade: true, solution: false -->
```python
""" [ASSERTS] Check the answer."""
print(b) # this should print: "[  0   1   4   9  16  25  36  49  64  81 100]"
assert b.shape == (11,)
```

Output:
```
[  0   1   4   9  16  25  36  49  64  81 100]
```

## [Task 3c] (5 Points) Create an 3x4 numpy array 'c' of random integers in [-5, +5[ (use numpy.random.randint).

<!-- grade_id: task3c-input, points: 0, grade: false, solution: true -->
```python
"""Put your code to create 'c' here.
"""
# YOUR CODE HERE
# raise NotImplementedError()
c = np.random.randint(-5, 5, size=(3, 4))
```

<!-- grade_id: task3c-test, points: 5, grade: true, solution: false -->
```python
""" [ASSERTS] Check the answer."""
print(c)
assert c.shape == (3,4)
```

Output:
```
[[ 0  0  4  3]
 [-4 -3 -3  3]
 [ 4 -2 -1  3]]
```

## [Task 3d] (5 Points) Let 'd' be the reshaped version of 'c' as a 6x2 array.

<!-- grade_id: task-3d-sol, points: 0, grade: false, solution: true -->
```python
"""Put your code to create 'd' here.
"""
# YOUR CODE HERE
# raise NotImplementedError()
d = c.reshape(6,2)
```

<!-- grade_id: task-3d-test, points: 5, grade: true, solution: false -->
```python
""" [ASSERTS] Check the answer."""
print(d)
assert d.shape == (6,2)
```

Output:
```
[[ 0  0]
 [ 4  3]
 [-4 -3]
 [-3  3]
 [ 4 -2]
 [-1  3]]
```

## [Task 3e] (5 Points) Use numpy or scipy to draw 1000 i.i.d. random samples from the Gaussian distribution with the mean given by the first digit of your UFID and the standard deviation given by the square root of the sum of all digits in your UFID. For example, if you UFID is 2801104 the mean would be 2 and the standard deviation is 4. Let 'e' be the array with the random samples.

<!-- grade_id: task3e-input, points: 0, grade: false, solution: true -->
```python
"""Put your code to create 'e' here.
"""
# YOUR CODE HERE
# raise NotImplementedError()
e = np.random.normal(loc=7, scale=6.4031, size=1000)
```

<!-- grade_id: task3e-test, points: 5, grade: true, solution: false -->
```python
""" [ASSERTS] Check the answer."""
assert e.shape == (1000,) or e.shape == (1000,1)
```

```python

```

```python

```